# <center>  Python for biologists
    
## Homework 9, Numpy

***Notion 0***. Feel free to start using shortcuts when working with Jupyter Notebook:
- To run the cell `Ctrl`+`Enter`
- To run the cell and step to the next one `Shift` + `Enter`
- To run the cell and create the new one `Alt` + `Enter`

There are two modes of action: cell-editing (the line on the left is green) and cell-selecting (the line on the left is blue). 

- cell-editing mode &#8594; `Esc` &#8594; cell-selecting mode
- cell-selecting mode &#8594; `Enter` &#8594; cell-editing mode

In cell-selecting mode you can:

- Delete  a cell - `DD`
- Undo deleting - `Z`

---

## Правила

- Это задание в первую очередь на знакомство с шириной и глубиной `numpy`. Откройте [документацию](https://numpy.org/doc/stable/reference/index.html#reference) и используйте её максимально при выполнении задания.
- Пишите код только там, где это от вас требуется (`# YOUR CODE HERE`)
- Написав, **удаляйте** `raise NotImplementedError() `
- Не переименовывайте переменные, если они уже созданы. Происходят автопроверки.
- Задания есть 2-х типов: пониженной (★☆) и повышенной (★★) сложности. Это нормально, если сложные задания кажутся сложными.
- Многие задания отмечены символом 🌐. Это значит что для его решения уже существует ОДНА готовая функция в библиотеке numpy. Осталось её только загуглить и подобрать параметры.
- Под заданием для вас приведены проверки с помощью `assert`. Они падают с ошибкой если решено не правильно и молчат если всё правильно.
- Перед отправкой проверьте что ваш ноутбук работает (юпитер-ноутбук). Нажмите на панели сверху `Kernel` -> `Restart Kernel and Run All Cells...`. Проверье что нет ошибок.
    

#### Погнали!

---
    
## I. Создание массивов

```python
import numpy as np
```

#### **Задание 1** (★☆) 🌐

Создайте вектор состоящий из 4-х нулей (массив `(4,)` ) 

<!-- grade_id: cell-ab35080a2e9048a6, points: 0, grade: false, solution: true -->
```python
result = np.zeros(4)
```

<!-- grade_id: cell-873f721b66636c4f, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([0, 0, 0, 0]))
```

#### **Задание 2** (★☆) 🌐

Создайте вектор состоящий из 4-х единиц (массив `(4,)` ) 

<!-- grade_id: cell-d768c26023eca51a, points: 0, grade: false, solution: true -->
```python
result = np.ones(4)
```

<!-- grade_id: cell-2b41e359eda9ca22, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([1, 1, 1, 1]))
```

#### **Задание 3** (★☆) 🌐

Создайте вектор состоящий из 10-и чисел от 11 до 20 включительно (массив `(10,)` ) 

<!-- grade_id: cell-6934c096c3e74b5c, points: 0, grade: false, solution: true -->
```python
result = np.arange(11,21)
```

<!-- grade_id: cell-ddfab69579023526, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([11, 12, 13, 14, 15, 16, 17, 18, 19, 20]))
```

#### **Задание 4** (★☆) 🌐

Выведите матрицу 2x3 заполненную нулями

<!-- grade_id: cell-668ba69fb6c560ba, points: 0, grade: false, solution: true -->
```python
result = np.zeros((2, 3))
```

<!-- grade_id: cell-be550edcb21a63b3, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[0, 0, 0], [0, 0, 0]]))
```

#### **Задание 4** (★☆) 🌐

Выведите количество элементов в массиве

<!-- grade_id: cell-f898b04d5add337c, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2]
                ]
)
result = mtr.size
```

<!-- grade_id: cell-a44d67999446b676, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, 6)
```

#### **Задание 5**(★☆) 🌐
 
Выведите размерность массива

<!-- grade_id: cell-67eef977740fff91, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2]
                ]
)
result = mtr.shape
```

<!-- grade_id: cell-c07cf06841b24e90, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, (2, 3))
```

#### **Задание 6** (★☆) 🌐

Создайте [единичную матрицу](https://ru.wikipedia.org/wiki/%D0%95%D0%B4%D0%B8%D0%BD%D0%B8%D1%87%D0%BD%D0%B0%D1%8F_%D0%BC%D0%B0%D1%82%D1%80%D0%B8%D1%86%D0%B0) 3x3

<!-- grade_id: cell-f907531ee4535156, points: 0, grade: false, solution: true -->
```python
result = np.eye(3)
#result = np.identity(3)
#or result = np.zeros((3, 3))
#result[0, 0] = 1
#result[1, 1] = 1
#result[2, 2] = 1
```

<!-- grade_id: cell-26974fc6eabd9367, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1., 0., 0.],
                                        [0., 1., 0.],
                                        [0., 0., 1.]]))
```

#### **Задание 7** (★☆) 🌐

Создайте вектор с 5-ю числами от 2 до 10 расположенными с равным интервальном между 2 и 10 (включительно)

<!-- grade_id: cell-ac65ec8f2cb1b4c2, points: 0, grade: false, solution: true -->
```python
result = np.arange(2, 11, 2)
```

<!-- grade_id: cell-f07bd80efe8c4683, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([ 2.,  4.,  6.,  8., 10.]))
```

#### **Задание 8** (★☆) 🌐

Создайте вектор с 3-мя числами вида $10^x$, где $x$ принимает значения между 2 и 6 (включительно). Здесь функция не та же, что и в предыдущем задании))

<!-- grade_id: cell-7b0bad4973d7cece, points: 0, grade: false, solution: true -->
```python
result = np.logspace(2, 6, 3)
```

<!-- grade_id: cell-abadae734e5ea344, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([1.e2, 1.e4, 1.e6]))
```

#### **Задание 9** (★★) 

Создайте матрицу 5х5 со значениями 3 в центре, 2 вокруг центра и 1 по краям

<img src='attachment:f534ff6f-19c5-4ee2-bdcf-8cf9eea0ef44.png' width=200>

<!-- grade_id: cell-f74eca2eece9e447, points: 0, grade: false, solution: true -->
```python
result = np.ones((5,5))
result[1:-1,1:-1] = 2
result[2,2] = 3
```

<!-- grade_id: cell-b16cb22b21275ceb, points: 7, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1., 1., 1., 1., 1.],
                                        [1., 2., 2., 2., 1.],
                                        [1., 2., 3., 2., 1.],
                                        [1., 2., 2., 2., 1.],
                                        [1., 1., 1., 1., 1.]]))
```

#### **Задание 10** (★★)

Создайте шахматную доску 8x8

<!-- grade_id: cell-79e9e4f0c86bdd16, points: 0, grade: false, solution: true -->
```python
result = np.tile(np.array([[0,1],[1,0]]), (4,4))
```

<!-- grade_id: cell-1b43bae7dea073bd, points: 7, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[0., 1., 0., 1., 0., 1., 0., 1.],
                                       [1., 0., 1., 0., 1., 0., 1., 0.],
                                       [0., 1., 0., 1., 0., 1., 0., 1.],
                                       [1., 0., 1., 0., 1., 0., 1., 0.],
                                       [0., 1., 0., 1., 0., 1., 0., 1.],
                                       [1., 0., 1., 0., 1., 0., 1., 0.],
                                       [0., 1., 0., 1., 0., 1., 0., 1.],
                                       [1., 0., 1., 0., 1., 0., 1., 0.]]))
```

## II. Операции над массивами

#### **Задание 11** (★☆)

Сложите две заданные матрицы

<!-- grade_id: cell-21cff1e910d36a96, points: 0, grade: false, solution: true -->
```python
mtr1 = np.array([[1, 1],
                [2, 2]
                ])
mtr2 = np.array([[0, 0],
                [3, 3]
                ])
result = mtr1 + mtr2
```

<!-- grade_id: cell-6626514d227b1c72, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1, 1],
                                        [5, 5]
                                        ]))
```

#### **Задание 12** (★☆)

Умножьте две заданные матрицы поэлементно

<!-- grade_id: cell-233aad32932a7432, points: 0, grade: false, solution: true -->
```python
mtr1 = np.array([[1, 1],
                [2, 2]
                ])
mtr2 = np.array([[0, 0],
                [3, 3]
                ])
result = mtr1*mtr2
```

<!-- grade_id: cell-4d374ad2be4104be, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[0, 0],
                                        [6, 6]
                                        ]))
```

#### **Задание 13** (★☆)

Умножьте две заданные матрицы матрично

<!-- grade_id: cell-cf37c10d511ad76a, points: 0, grade: false, solution: true -->
```python
mtr1 = np.array([[1, 1],
                [2, 2]
                ])
mtr2 = np.array([[0, 0],
                [3, 3]
                ])
result = mtr1 @ mtr2
#np.matmul(mtr1, mtr2)
#np.dot(mtr1, mtr2)
```

<!-- grade_id: cell-e892bbd6ff7ee7db, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[3, 3],
                                        [6, 6]
                                        ]))
```

#### **Задание 14** (★☆) 🌐

Вам дана матрица 2x3, переведите её в матрицу 3x2

<!-- grade_id: cell-097bd3d452be27a8, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2]
                ]
)
result = mtr.reshape(3,2)
```

<!-- grade_id: cell-e2699ce9b7be179f, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1, 1],
                                        [1, 2],
                                        [2, 2]]))
```

#### **Задание 15** (★☆) 🌐

Вам дана матрица 3x3, переведите её в вектор из 9 элементов

<!-- grade_id: cell-22ca17c2d42885e5, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2],
                [3, 3, 3]
                ]
)
result = mtr.reshape(-1)
```

<!-- grade_id: cell-2d9eea5a9ad54ea7, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([1, 1, 1, 2, 2, 2, 3, 3, 3]))
```

#### **Задание 16** (★☆) 🌐

Вам дана матрица 3x3, транспонируйте её

<!-- grade_id: cell-992989b3092b93cc, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2],
                [3, 3, 3]
                ]
)
result = mtr.T
#np.transpose(mtr)
```

<!-- grade_id: cell-cb03c5b94f8ba89d, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1, 2, 3],
                                        [1, 2, 3],
                                        [1, 2, 3]]))
```

#### **Задание 17** (★☆) 🌐

Вам дана матрица 3x3, поверните её на 270 градусов (3 раза по 90 градусов)

<!-- grade_id: cell-19c8f0e83cdab328, points: 0, grade: false, solution: true -->
```python
mtr = np.array([
                [1, 1, 1],
                [2, 2, 2],
                [3, 3, 3]
                ]
)
result = np.rot90(mtr, 3)
```

<!-- grade_id: cell-b6da2013fe4d6c9d, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[3, 2, 1],
                                        [3, 2, 1],
                                        [3, 2, 1]]))
```

#### **Задание 18** (★☆) 🌐
 
Вам дано две матрицы 2x2, получите матрицу 4x2 "поставив" эти матрица одна на другую (первая сверху).

<!-- grade_id: cell-de7def3eef99b91c, points: 0, grade: false, solution: true -->
```python
mtr1 = np.array([[1, 1],
                [2, 2]
                ])
mtr2 = np.array([[0, 0],
                [3, 3]
                ])
result = np.concatenate((mtr1, mtr2), axis=0)
```

<!-- grade_id: cell-b1c48c547f14ce2c, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[1, 1],
                                        [2, 2],
                                        [0, 0],
                                        [3, 3]]))
```

#### **Задание 19** (★★)

Вам дано 10 пар декартовых координат точек на плоскости (матрица 5x2). Переведите все пары декартовых координат $(x, y)$ в [полряные координаты](https://www.fxyz.ru/%D1%84%D0%BE%D1%80%D0%BC%D1%83%D0%BB%D1%8B_%D0%BF%D0%BE_%D0%BC%D0%B0%D1%82%D0%B5%D0%BC%D0%B0%D1%82%D0%B8%D0%BA%D0%B5/%D1%84%D1%83%D0%BD%D0%BA%D1%86%D0%B8%D0%B8_%D0%B8_%D0%B3%D1%80%D0%B0%D1%84%D0%B8%D0%BA%D0%B8/%D0%BA%D0%BE%D0%BE%D1%80%D0%B4%D0%B8%D0%BD%D0%B0%D1%82%D1%8B/%D0%BF%D0%B5%D1%80%D0%B5%D0%B2%D0%BE%D0%B4_%D0%B4%D0%B5%D0%BA%D0%B0%D1%80%D1%82%D0%BE%D0%B2%D1%8B%D1%85_%D0%BA%D0%BE%D0%BE%D1%80%D0%B4%D0%B8%D0%BD%D0%B0%D1%82_%D0%B2_%D0%BF%D0%BE%D0%BB%D1%8F%D1%80%D0%BD%D1%8B%D0%B5/) $(\rho, \varphi$). У вас также должна получиться матрица 10x2.

<!-- grade_id: cell-84a629e3a0621a90, points: 0, grade: false, solution: true -->
```python
coords = np.array([[0.65, 0.44],
              [0.96, 0.08],
              [0.16, 0.11],
              [0.36, 0.07],
              [0.89, 0.78],
              [0.11, 0.97],
              [0.08, 0.24],
              [0.33, 0.61],
              [0.86, 0.31],
              [0.98, 0.61]])
x,y = coords[:,0], coords[:,1]
a = np.hypot(x, y)
b = np.arctan2(y,x)
result = np.array([a, b]).T
```

<!-- grade_id: cell-b283ca9679052543, points: 9, grade: true, solution: false -->
```python
assert np.allclose(result, np.array([[0.78492038, 0.59506962],
                                     [0.96332757, 0.08314123],
                                     [0.19416488, 0.60228735],
                                     [0.36674242, 0.19204802],
                                     [1.18342723, 0.71962491],
                                     [0.97621719, 1.45787667],
                                     [0.25298221, 1.24904577],
                                     [0.69354164, 1.07490184],
                                     [0.91416629, 0.34596727],
                                     [1.15433964, 0.55676277]]))
```

#### **Задание 20** (★★)

Вам дано два вектора $x$, $y$, посчитайте определитель [матрицы Коши](https://ru.wikipedia.org/wiki/%D0%9C%D0%B0%D1%82%D1%80%D0%B8%D1%86%D0%B0_%D0%9A%D0%BE%D1%88%D0%B8_(%D0%BB%D0%B8%D0%BD%D0%B5%D0%B9%D0%BD%D0%B0%D1%8F_%D0%B0%D0%BB%D0%B3%D0%B5%D0%B1%D1%80%D0%B0)) $det A$, где элементы $A$ вычисляются по формуле $a_{ij} = \frac{1}{x_i - y_j} $

<!-- grade_id: cell-1f52f323aef2aa20, points: 0, grade: false, solution: true -->
```python
x = np.array([0, 1, 2, 3, 4, 5, 6, 7])
y = np.array([ 2.5,  3.5,  4.5,  5.5,  6.5,  7.5,  8.5,  9.5])

result = np.linalg.det(1 / np.subtract.outer(x, y))
```

<!-- grade_id: cell-b4d3a3abf081a348, points: 6, grade: true, solution: false -->
```python
assert np.isclose(result, -0.0020386723097196844)
```

#### **Задание 21** (★☆)

Вам дан вектор значений, посчитайте [сигмоиду](https://en.wikipedia.org/wiki/Sigmoid_function) от этого вектора

<!-- grade_id: cell-fedb26e52cd43fd0, points: 0, grade: false, solution: true -->
```python
x = np.array([0, 1, 2, 3, 4, 5, 6, 7])

result = result = 1 / ( 1 + np.exp(-x))
```

<!-- grade_id: cell-8533270e0383f3cb, points: 2, grade: true, solution: false -->
```python
assert np.allclose(result, np.array([0.5, 0.73105858, 0.88079708, 0.95257413, 0.98201379, 0.99330715, 0.99752738, 0.99908895]))
```

## III. Поиск в массивах

#### **Задание 22** (★☆)

Вам дан вектор значений, сделайте маску для поиска положительных чисел в массиве (строго больше, в ответе просто маска)

<!-- grade_id: cell-3ecfd1321ac38bcd, points: 0, grade: false, solution: true -->
```python
arr = np.array([-1, 2, 0, -10, 5, 6, 7])

result = arr > 0
```

<!-- grade_id: cell-e3f6d04510d24f0a, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([False,  True, False, False,  True,  True,  True]))
```

#### **Задание 23** (★☆) 🌐

Вам дан вектор, получите индексы ненулевых элементов в векторе

<!-- grade_id: cell-e4ee9e8833db8a72, points: 0, grade: false, solution: true -->
```python
arr = np.array([-1, 2, 0, -10, 5, 6, 7])

result = np.flatnonzero(arr)
```

<!-- grade_id: cell-63e39120fc1b281d, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([0, 1, 3, 4, 5, 6]))
```

#### **Задание 24** (★★) 🌐

Вам дан массив, выведите массив где все отрицательные числа остались на том же месте, а на месте неотрицательных чисел стоят их квадраты.

<!-- grade_id: cell-304be7d582634d49, points: 0, grade: false, solution: true -->
```python
arr = np.array([
               [-1, 2, 0],
               [-10, 5, 6],
               [11, -11, -0.5]
               ])

result = np.where(arr > 0, arr ** 2, arr)
print(result)
```

<!-- grade_id: cell-1cb68478de452739, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[ -1. ,   4. ,   0. ],
                                        [-10. ,  25. ,  36. ],
                                        [121. , -11. ,  -0.5]]))
```

#### **Задание 25** (★☆) 🌐

Вам дан вектор значений и вектор индексов. Выберете значения соответствующие индексам.

<!-- grade_id: cell-1759eabfe0177f5e, points: 0, grade: false, solution: true -->
```python
arr = np.array([-1, 2, 0, -10, 5, 6, 7])
indices = [0, 1, -1]

result = arr.flat[indices]
```

<!-- grade_id: cell-08a9f8644fe4c62a, points: 3, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([-1,  2,  7]))
```

#### **Задание 26** (★★) 🌐
 
Вам дан вектор значений и функция. Подберите такую переменную `indexes`, чтобы получить 3-мерный массив, где на главной диагонали везде находится последнее значение из вектора, а в остальных ячейках - первое.

<!-- grade_id: cell-4aca204d0d46dba6, points: 0, grade: false, solution: true -->
```python
arr = np.array([7, 0, -10, 5, 6, 2])

indices = np.where(np.eye(3), -1, 0)
result = np.take(arr, indices)
```

<!-- grade_id: cell-31db76215e8df45f, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([[2, 7, 7],
                                        [7, 2, 7],
                                        [7, 7, 2]]))
```

#### **Задание 27** (★☆) 🌐

Верните индексы максимального элемента в каждом столбце матрицы. 

<!-- grade_id: cell-2628abb3dd278812, points: 0, grade: false, solution: true -->
```python
arr = np.array([
               [-1, 2, 0],
               [-10, 5, 6]
               ])


result = np.argmax(arr, axis=0)
```

<!-- grade_id: cell-d4c6da3a97935c2c, points: 2, grade: true, solution: false -->
```python
assert np.array_equal(result, np.array([0, 1, 1]))
```

## IV. Великий рандом

> Отдельная важная часть библиотеки `numpy` - работа со случайными числами. Мы не успели посмотреть её на лекции, но тут ничего сложного. Для лучшего понимания советую почитать про [Генератор псевдослучайных чисел](https://ru.wikipedia.org/wiki/%D0%93%D0%B5%D0%BD%D0%B5%D1%80%D0%B0%D1%82%D0%BE%D1%80_%D0%BF%D1%81%D0%B5%D0%B2%D0%B4%D0%BE%D1%81%D0%BB%D1%83%D1%87%D0%B0%D0%B9%D0%BD%D1%8B%D1%85_%D1%87%D0%B8%D1%81%D0%B5%D0%BB) (на котором основывается `numpy` и не только `numpy`) и wiki `numpy` про модуль [Random](https://numpy.org/doc/stable/reference/random/index.html).

В этой части используйте самый простой функционал - через `np.random.[функция]`.  **НЕ НАДО** создавать отдельные генераторы через `RandomState`, `Generator`, `default_rng`, `BitGenerator` и т.п.

#### **Задание 28** (★☆) 🌐

Зафиксируйте сид ГПСЧ numpy числом `3_099_734_149`.

<!-- grade_id: cell-c65294d179583d1b, points: 0, grade: false, solution: true -->
```python
result = np.random.seed(3_099_734_149)
```

<!-- grade_id: cell-2f9b56bcfbed2b42, points: 2, grade: true, solution: false -->
```python
assert np.random.get_state()[1][0] == 3_099_734_149
```

#### **Задание 29** (★☆) 

Напишите **РОВНО** 3 или 4 предложения на русском, где вашими словами объясняется что такое `seed` и зачем он нужен.

<!-- grade_id: cell-c347d894f44b6e89, points: 1, grade: true, solution: true -->
seed - это определенное число, которое задает начальное состояние для генератора чисел. Он нужен для того, чтобы результаты, запускаемые на разных устройствах/в разное время/разными людьми были воспроизводимы, т.е. были одинаковыми при первом запуске. Благодаря seed задаваемые случайные числа становятся предсказуемыми.

> А также подумайте, если мы зафиксировали сид, то нам будут выпадать каждый раз одни и те же случайные значения? Или нет? Что в данном случае означает слово "воспроизводимость". Ответ: каждый раз генерируя случайное значение, вы будете получать новое значение. Но если вы перезапустите ядро, то при повторении процедуры вы будете получать те же самые случайные значения в том же порядке. Почему так? (ответьте в своей голове)

#### **Задание 30** (★☆) 🌐

Вам предложен список оценок за курс по питону. Выберите случайную оценку из списка.

<!-- grade_id: cell-2bcbeb9b3d41641b, points: 0, grade: false, solution: true -->
```python
marks = ['A', 'B', 'C', 'D', 'E', 'F']

mark = np.random.choice(marks)
```

<!-- grade_id: cell-34d25c5ef3bc98ad, points: 2, grade: true, solution: false -->
```python
assert mark == 'C'
```

#### **Задание 31** (★☆) 🌐

Сгенерируйте случайное целое число между 0 (inclusive) и 10 (exclusive)

<!-- grade_id: cell-de4243d34bd4ee52, points: 0, grade: false, solution: true -->
```python
number = np.random.randint(0, high=10)
```

<!-- grade_id: cell-9931b1c33d71c91d, points: 2, grade: true, solution: false -->
```python
assert number == 6
```

#### **Задание 32** (★☆) 🌐

Сгенерируйте массив [вероятностей](https://ru.wikipedia.org/wiki/%D0%92%D0%B5%D1%80%D0%BE%D1%8F%D1%82%D0%BD%D0%BE%D1%81%D1%82%D1%8C#%D0%A1%D0%B2%D0%BE%D0%B9%D1%81%D1%82%D0%B2%D0%B0_%D0%B2%D0%B5%D1%80%D0%BE%D1%8F%D1%82%D0%BD%D0%BE%D1%81%D1%82%D0%B8)  размерностью 2x2

<!-- grade_id: cell-59334c907c2af21e, points: 0, grade: false, solution: true -->
```python
probs = np.random.rand(2, 2)
```

<!-- grade_id: cell-a37548867423bd70, points: 2, grade: true, solution: false -->
```python
assert np.allclose(probs, np.array([[0.35928318, 0.5584223 ],
                                 [0.12939988, 0.68296292]])
                  )
```

#### **Задание 33** (★☆) 🌐

Сгенерируйте 5 случайных чисел из равномерного распределения с параметрами $a=2, b=55$

<!-- grade_id: cell-231edc4150c82d07, points: 0, grade: false, solution: true -->
```python
numbers = np.random.uniform(low=2.0, high=55.0, size=5)
```

<!-- grade_id: cell-ddbbb08c115bf28e, points: 2, grade: true, solution: false -->
```python
assert np.allclose(numbers, np.array([21.28655049, 54.23020856, 41.54699812, 31.22932382, 12.25757421]))
```

#### **Задание 34** (★☆) 🌐

Сгенерируйте 5 случайных чисел из нормального распределения с параметрами $\mu = 10, \sigma = 15$

<!-- grade_id: cell-d219897f085f07ff, points: 0, grade: false, solution: true -->
```python
mu, sigma = 10, 15
numbers = np.random.normal(mu, sigma, 5)
```

<!-- grade_id: cell-3699a882c0bdb6e3, points: 2, grade: true, solution: false -->
```python
assert np.allclose(numbers, np.array([ 0.80725362, 20.30002062,  9.92106564,  3.60193382,  8.64049947]))
```

#### **Задание 35** (★☆) 🌐

Вам дан вектор объектов. Сделайте случайную перестановку в этом векторе.

<!-- grade_id: cell-8c4bdf4c8d1ca806, points: 0, grade: false, solution: true -->
```python
str_ = list('биоинформатик')

perm_str = np.random.permutation(str_)

perm_str = ''.join(perm_str)
```

<!-- grade_id: cell-885b64fb5f179af8, points: 2, grade: true, solution: false -->
```python
assert perm_str == 'крибамофоинит'
```

### **Задание XXX**

Нажмите на панели сверху `Kernel` -> `Restart Kernel and Run All Cells...`. Проверье что нет ошибок.

---

Надеюсь, у тебя всё получилось! Но в любом случае - ты молодец:)🥳

---